# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(L):
    '''Takes a string as input and creates a new list with the length of items in the new list. Finally finds the max value and position of the max value in the new list. Return the (first) position of the max value'''
    length = []
    for i in L:
        length.append(int(len(i)))
    return length.index(max(length))
        

### END SOLUTION
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

Output:
```
1
```

```python
def longest(L):
    """(list) -> int
    Finds the longest string in this list.
    
    >>> longest(['my','name','is','Paul Boal'])
    3
    
    >>> longest("happy birthday to me".split(" "))
    1
    """
    
    position = 0
    length = 0
    
    for pos, word in enumerate(L):
        if len(word) > length:
            position = pos
            length = len(word)
    
    return position
    
    
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['my','name','is','Paul Boal'])
Expecting:
    3
ok
Trying:
    longest("happy birthday to me".split(" "))
Expecting:
    1
ok
```

### 25.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION

def f_to_c(temp_f):
    """(float) -> float
    Converts Farenheit to Celsius
    
    >>> f_to_c(212)
    100.0
    
    >>> f_to_c(32)
    0.0
    
    >>> f_to_c(98.6)
    37.0
    """
    
    return (5/9 * (temp_f - 32))

### END SOLUTION
```

```python
### BEGIN SOLUTION
def f_to_c(F):
    """(float) -> float
    
    >>> f_to_c(212)
    100.0
        
    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0
    """
    return (5/9) * (F-32)

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
f_to_c(212)
```

Output:
```
100.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
def los(admit, discharge, diagnosis):
    """(date, date, str) -> int
    Return length of stay in days based on the rules provided.
    If diagnosis is "Observation" then LOS = 0
    Otherwise LOS is the number of whole days between admin and discharge dates.
    
    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5
    
    >>> los(date(2020,1,5), date(2020,1,8), "Observation")
    0
    """
    d = (discharge - admit).days + 1
    # Adding one to the days counted, if admited the minimum days should read 1 and include the first day.
    
    if diagnosis == 'Observation':
        return 0
    else:
        return d
    
### END SOLUTION
```

```python
def los(admit, discharge, diagnosis):
    """(date, str, date)->int
    Computes length of stay in days.  If diagnosis is "Observation" then set the length of stay to 0.

    >>> los(date(2021,1,3), date(2021,1,8), 'COPD')
    5
    
    >>> los(date(2021,1,5), date(2021,1,9), 'Hypertension')
    4
    
    >>> los(date(2021,1,23), date(2021,1,25), 'Observation')
    0
    """
    if diagnosis == "Observation":
        return 0
    else:
        return (discharge - admit).days
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2021,1,3), date(2021,1,8), 'COPD')
Expecting:
    5
ok
Trying:
    los(date(2021,1,5), date(2021,1,9), 'Hypertension')
Expecting:
    4
ok
Trying:
    los(date(2021,1,23), date(2021,1,25), 'Observation')
Expecting:
    0
ok
```

```python
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for e, v in encounters.items():
    print(los(v['admit'],v['discharge'],v['diagnosis']))
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
from datetime import date
test1 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test2 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
        "E4234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,7) }
}

### BEGIN SOLUTION
def average_los(encounters):
    """
    (dict) -> float
    Return average length of stay for the encounters.
    
    >>> average_los(test1)
    5
    """
    enc = encounters.keys()
    los_list=[]
    for e, v in encounters.items():
        los_list.append(los(v['admit'],v['discharge'],v['diagnosis']))
    
    
    
    return sum(los_list)/len(los_list)

average_los(test2)
### END SOLUTION
```

Output:
```
<ipython-input-2-e6ce9781b2ce> in <module>
     27     return sum(los_list)/len(los_list)
     28 
---> 29 average_los(test2)
     30 ### END SOLUTION

<ipython-input-2-e6ce9781b2ce> in average_los(encounters)
     21     los_list=[]
     22     for e, v in encounters.items():
---> 23         los_list.append(los(v['admit'],v['discharge'],v['diagnosis']))
     24 
     25 

NameError: name 'los' is not defined
```

```python
from datetime import date
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(average_los(encounters)) == 2.75
# 2.75*5 = 13.75, so 2.75 cannot be the actual correct answer, considering [6,5,2,0,1] the average will be 2.8
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-16-98eace87cc6d> in <module>
      8 }
      9 
---> 10 assert(average_los(encounters)) == 2.75
     11 # 2.75*5 = 13.75, so 2.75 cannot be the actual correct answer, considering [6,5,2,0,1] the average will be 2.8

AssertionError:
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION

def c_to_f(centigrade):
    '''
    (float) -> float
    Returns the Farenheit value of temperature
    >>> c_to_f(100)
    212.0
    
    >>> c_to_f(0)
    32.0
    
    >>> c_to_f(37)
    98.6
    
    '''
    return (centigrade *(9/5))+32

### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

---

```python
%%bash
git pull
git add week05_assignment_2.ipynb
git commit -a -m "Submitting the week 5 programming assignment"
git push
```

Output:
```
Already up to date.
[main 27f31cb] Submitting the week 5 programming assignment
 1 file changed, 696 insertions(+)
 create mode 100644 week05/week05_assignment_2.ipynb
To github.com:Rathnareddy163/hds5210-2021.git
   66ed534..27f31cb  main -> main
```


---

If the message above says something like _Submitting the week 5 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.

```python

```

```python

```